---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

You should complete this exercise ASAP

`@instructions`
Do it again and again, if need any help please contact to CR

`@hint`
it's easy, just practice. Practice makes main perfect

`@pre_exercise_code`
```{python}
import numpy as np
np.array([1,2,3,4])
np.shape()
```

`@sample_code`
```{python}
import numpy as np
np.array([1,2,3,4])
np.shape()
```

`@solution`
```{python}
import numpy as np
np.array([1,2,3,4])
np.shape()
```

`@sct`
```{python}
You are right
```
